![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
import matplotlib.pyplot as plt
```

```python
from itertools import combinations
```

# Interpolação em pontos diferentes

Vamos continuar observando a interpolação de Lagrange.
Agora, vamos escolher pontos que não são uniformemente espaçados nos intervalos.

<!-- grade_id: cell-c706d1ad462d25b8, points: 0, grade: false, solution: true -->
```python
def lagrange_coefs(xs,ys):
    """ 
    Coeficientes (em ordem crescente) do polinômio interpolador dos pontos (x_i, y_i).
    """
    assert(len(xs) == len(ys))
    # YOUR CODE HERE
    raise NotImplementedError()

def lagrange(xs,ys):
    """ 
    Retorna uma função que calcula polinômio interpolador de Lagrange dos pontos (x_i, y_i). 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

# Parte 1: Escolhendo três pontos para interpolar

Vamos usar 11 pontos no intervalo $[-1,1]$, e ver qual escolha de pontos dá o menor erro de interpolação.

## Questão 1: Medindo o máximo do erro

Para observar o erro entre $f$ e o (um) polinômio interpolador $P_f$, precisamos de uma função para medir o erro.
Vamos usar o máximo do erro no intervalo $[a,b]$.
Aqui, novamente, vamos usar os "pontos médios" de uma subdivisão,
para evitar privilegiar os extremos do intervalo.

<!-- grade_id: cell-2d60ad4002a4ba00, points: 0, grade: false, solution: true -->
```python
def maxerr(f, P, a, b, n=100):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-d1ebcf4b35edf552, points: 1, grade: true, solution: false -->
```python
assert maxerr(np.sin, np.cos, 0, np.pi) > 1.41
```

<!-- grade_id: cell-540b328661601784, points: 1, grade: true, solution: false -->
```python
assert maxerr(np.sin, np.cos, 16, 17, n=300) < 0.7
```

## Questão 2: Variando os pontos

Considere os 11 pontos abaixo.
Para todas as $\binom{11}{3} = 165$ escolhas possíveis de 3 pontos,
calcule o máximo do erro de interpolação (no intervalo $[-1,1]$)
entre o polinômio e a função `np.exp`.

<!-- grade_id: cell-4f250181db6daaf5, points: 2, grade: true, solution: true -->
```python
ts = np.linspace(-1,1,11)
errs_max = []

for xs in combinations(ts, 3):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Abaixo, vamos ver como ficam os erros, tanto em escala natural como em escala log.

```python
def histogramas(errs):
    fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))
    ax1.hist(errs, bins=20)
    ax2.hist(np.log2(errs), bins=20)

    ax1.set_title("Distribuição dos erros")
    ax1.set_xlabel("Erro")
    ax2.set_title("Distribuição dos erros")
    ax2.set_xlabel("Log em base 2 do erro")
```

```python
histogramas(errs_max)
```

O que você percebe quanto à distribuição dos erros?

<!-- grade_id: cell-caca0bbed48c107d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Qual a combinação de pontos teve o maior erro? E qual teve o menor erro?

Dica: Você pode usar as funções `np.argmax()`, `np.argmin()`.

<!-- grade_id: cell-a60cf8dfb94f60ef, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Faça o gráfico dos dois polinômios interpoladores correspondentes,
sobrepostos ao gráfico da função exponencial.

<!-- grade_id: cell-e6e8b1bda846615a, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-7ca442f203faa0d8, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

Comente, por exemplo:
- O comportamento destas curvas é esperado, dados os pontos usados para interpolação?
- Quão grandes estão os erros?

<!-- grade_id: cell-a010bbf5a478fde2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 3: Pontos uniformemente espaçados

Para comparar com o polinômio obtido por pontos igualmente espaçados,
modifique a função `histogramas` para receber outro argumento,
e traçar uma linha vertical na posição correspondente
(nos dois eixos!).

<!-- grade_id: cell-abd2e4b7210dc5ca, points: 1, grade: true, solution: true -->
```python
def histogramas(errs, y_destaque):
    fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))
    ax1.hist(errs, bins=20)
    ax2.hist(np.log2(errs), bins=20)

    ax1.set_title("Distribuição dos erros")
    ax1.set_xlabel("Erro")
    ax2.set_title("Distribuição dos erros")
    ax2.set_xlabel("Log em base 2 do erro")
    # YOUR CODE HERE
    raise NotImplementedError()
```

Agora, calcule o erro máximo cometido pelo polinômio interpolador com 3 pontos igualmente espaçados em $[-1,1]$:

<!-- grade_id: cell-08ef1b4c9daeb61d, points: 1, grade: true, solution: true -->
```python
# Dê sua resposta da forma
err_unif = 0
# YOUR CODE HERE
raise NotImplementedError()
```

```python
histogramas(errs_max, err_unif)
```

Comente.

<!-- grade_id: cell-e7b48c534994c65e, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 4: Mais pontos!

Agora, repita o estudo anterior para 31 pontos uniformemente espaçados em $[-1,1]$.

<!-- grade_id: cell-a14d9c1ee59082b9, points: 0, grade: false, solution: true -->
```python
def all_errs(f, n, pts, errfun):
    """Retorna todos os erros dos polinômios interpoladores para a função  f,
    gerados pela escolha de  n  pontos da lista/array  pts.

    A função  errfun  recebe duas funções  f e g  e retorna o erro entre elas."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-837ec20ed5dfb3e2, points: 1, grade: true, solution: false -->
```python
def aux_maxerr(f, g):
    return maxerr(f, g, -1, 1, n=200)

errs31 = all_errs(np.exp, 3, np.linspace(-1,1,31), aux_maxerr)
assert len(errs31) == 31*30*29//6
```

<!-- grade_id: cell-65b38e1aa0aa06bc, points: 1, grade: true, solution: false -->
```python
assert np.abs(min(errs31) - 0.05) < 1e-3
assert 1.9 < max(errs31) < 2.0
```

Vejamos os histogramas dos erros:

```python
histogramas(errs31, err_unif)
```

O que você observa, em comparação com o caso inicial com apenas 11 pontos?

<!-- grade_id: cell-a9646523b66a28cb, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Calcule o melhor e pior polinômios:

<!-- grade_id: cell-3a962475a27f675e, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

ax = plt.gca()
```

<!-- grade_id: cell-df2b40c31748029f, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
```

Faça, agora, o gráfico do erro para os dois polinômios.

<!-- grade_id: cell-c2228efe7331689c, points: 1, grade: true, solution: true -->
```python
fig, (ax1, ax2) = plt.subplots(ncols=2, figsize=(13,4))

# YOUR CODE HERE
raise NotImplementedError()

ax1.grid()
ax2.grid()
```

<!-- grade_id: cell-0553a136bf74a232, points: 1, grade: true, solution: false -->
```python
for ax in (ax1, ax2):
    assert len(ax.lines) == 1
    assert ax.get_title() != ""
```

## Questão 5: Outras funções

Repita a análise com 31 pontos (ou mais, se seu código e computador forem rápidos o suficiente!)
para as seguintes funções:

- cosseno, em $[-1,1]$
- log, em $[2,10]$,
- $\frac{1}{1 + x^2}$ em $[-1,1]$

Para evitar repetir muito código, crie uma função `analise` para fazer as contas e os gráficos!

<!-- grade_id: cell-eecb4f2f6c9b2310, points: 2, grade: true, solution: true -->
```python
def analise(f, a, b, deg, npts, nfine):
    """Análise de erro de interpolação da função  f, no intervalo  [a,b], com polinômios de grau  deg.
    Usa  npts  pontos igualmente espaçados para escolher os nós de interpolação,
    e  nfine  pontos para calcular erros / fazer gráficos.
    
    Faz
    - os histogramas dos erros
    - gráficos dos erros do melhor / pior polinômio
    E retorna os pontos correspondentes à melhor / pior interpolação."""
    assert npts > deg
    tfine = np.linspace(a,b,nfine)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# Verifique que está dando certo!
analise(np.exp, -1, 1, 2, 31, 200)
```

```python
analise(np.cos, -1, 1, 2, 31, 200)
```

```python
analise(np.log, 2, 10, 2, 31, 200)
```

```python
def cauchy(x):
    x = np.asarray(x)
    return 1/(1 + x**2)
```

```python
analise(cauchy, -1, 1, 2, 31, 200)
```

Comente:
- como ficaram as distribuições dos erros
- como ficam as curvas de erro para o melhor / pior caso
    - se isso faz sentido

<!-- grade_id: cell-20090005accca597, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, comente:
- quais foram os pontos escolhidos para melhor / pior caso
- e novamente se isso faz sentido.

<!-- grade_id: cell-49ab92e4266a7a31, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

# Parte 2: Regressão

Em vez de tentarmos encontrar o melhor polinômio de grau $2$ que aproxima $f$
percorrendo todas as $\binom{N}{3}$ combinações possíveis de pontos igualmente espaçados,
vamos fazer algo diferente:
- Vamos calcular a função $f$ nos $N$ pontos; e
- ajustar um polinômio de grau $2$ por regressão.

## Questão 6: Calculando o ajuste

Escreva uma função `poly_reg`, que recebe $f$, $[a,b]$, o número de pontos $N$ e o grau desejado $d$,
e retorna o melhor polinômio de grau $d$ por mínimos quadrados.

<!-- grade_id: cell-f68b0b0da88cbf15, points: 0, grade: false, solution: true -->
```python
def poly_reg(f, a, b, N, d):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2bf26a5fb1b2d302, points: 2, grade: true, solution: false -->
```python
P_reg = poly_reg(np.sin, -1, 1, 31, 2)
assert np.isclose(P_reg(0), 0)
assert np.isclose(P_reg(1), - P_reg(-1))
```

<!-- grade_id: cell-497b6770f04d71ee, points: 1, grade: true, solution: false -->
```python
P_reg = poly_reg(cauchy, -1, 1, 31, 2)
assert maxerr(cauchy, P_reg, -1, 1, 100) < 0.06
```

<!-- grade_id: cell-d97d2bfd73d1d4ff, points: 1, grade: true, solution: false -->
```python
P_reg = poly_reg(cauchy, -1, 1, 31, 20)
assert 1e-7 < maxerr(cauchy, P_reg, -1, 1, 2000) < 5e-7
```

<!-- grade_id: cell-55954f12f4f8a5e8, points: 1, grade: true, solution: false -->
```python
P_reg = poly_reg(cauchy, -1, 1, 301, 20)
assert 5e-9 < maxerr(cauchy, P_reg, -1, 1, 20000) < 20e-9
```

Comente os três últimos `assert`s:
- o que você deduz de aumentar o grau? e aumentar o número de pontos?
- porque usar 2000 e 20000 pontos para o máximo do erro?

<!-- grade_id: cell-5e69e9313de954c4, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
P_reg = poly_reg(cauchy, -1, 1, 3001, 20)
maxerr(cauchy, P_reg, -1, 1, 20000)
```

O que esta última caixa sugere?

<!-- grade_id: cell-7193e8a89fbf5f31, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 7: comparações

Agora, compare os erros para:
- o melhor polinômio de grau 2, buscando em 31 pontos;
- o melhor polinômio de grau 2, fazendo regressão em 31 pontos;
- o polinômio de grau 2, usando pontos igualmente espaçados.

<!-- grade_id: cell-d6d38df4baf54481, points: 2, grade: true, solution: true -->
```python
def compara_erros(f, a, b, npts, deg, nfine):
    """Cria três gráficos dos erros dos polinômios de grau  deg, usando as 3 estratégias acima,
    para  f  no intervalo [a,b], a partir de  npts  pontos."""
    assert npts > deg
    tfine = np.linspace(a,b,nfine)
    fig, (ax1, ax2, ax3) = plt.subplots(ncols=3, figsize=(18,4), sharey=True)
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
compara_erros(np.exp, -1, 1, 31, 2, 200)
```

```python
compara_erros(np.cos, -1, 1, 31, 2, 200)
```

O que você percebe nestes gráficos?

<!-- grade_id: cell-2993fa62eb1c9e28, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

```python
compara_erros(np.log, 2, 10, 31, 2, 200)
compara_erros(cauchy, -1, 1, 31, 2, 200)
```

Para estas funções, o que você observa de similar / diferente quanto às duas anteriores?

<!-- grade_id: cell-cca30a1a0c9cab4d, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 8: Aumentando o grau

Vamos agora tentar achar polinômios de grau 4.
A caixa abaixo pode levar um certo tempo para rodar... (tente ficar abaixo de 1 min, mas não é o mais importante)

```python
%time compara_erros(cauchy, -1, 1, 31, 4, 200)
```

O que você percebe de similar / diferente neste caso, com relação aos polinômios de grau 2?

<!-- grade_id: cell-fc1c2e2af0bf63b2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Qual é a parte mais demorada de `compara_erros`?
Porquê?

<!-- grade_id: cell-4160689bf7ce45fb, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 9: Aumentando ainda mais o grau

Uma outra forma de calcular aproximações é usar o polinômio de Taylor.
Defina, abaixo, a função `exp_taylor` que calcula a aproximação da exponencial
pelo polinômio de Taylor:
$$ e^x \simeq 1 + x + \frac{x^2}{2} + \ldots + \frac{x^n}{n!}. $$

<!-- grade_id: cell-b0b0efaf037a64a4, points: 0, grade: false, solution: true -->
```python
def exp_taylor(xs, deg):
    """Série de Taylor da exponencial, com grau  deg"""
    xs = np.asarray(xs, dtype=np.float)
    cur = np.ones_like(xs)
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8c0c1bc36ce6c98b, points: 1, grade: true, solution: false -->
```python
assert exp_taylor(1, 3) == 1 + 1 + 1/2 + 1/6
```

<!-- grade_id: cell-16b51b53b35ddcd6, points: 1, grade: true, solution: false -->
```python
xs = [-1,1,0.6]
assert not np.allclose(exp_taylor(xs, 5), np.exp(xs))
assert np.allclose(exp_taylor(xs, 8), np.exp(xs))
```

Agora, faça um gráfico comparando os erros, no intervalo $[-1,1]$:
- do polinômio de Taylor de grau 5 para a exponencial
- do polinômio interpolador de grau 5
- do polinômio de grau 5 obtido por regressão em 100 pontos.

<!-- grade_id: cell-e9bb62be17a7f1f8, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

plt.legend(title="Método", bbox_to_anchor=(1, 0.5), loc="center left")
plt.grid()
ax = plt.gca()
```

<!-- grade_id: cell-96e053bf343b1da8, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 3
assert len(ax.legend().texts) == 3
assert ax.get_title() != ""
```

O que você observa quanto aos erros de cada método?

<!-- grade_id: cell-41ac6eabe481c9c2, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

Agora, faça para grau 10

<!-- grade_id: cell-f8afda04e91b2cc4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

plt.legend(title="Método", bbox_to_anchor=(1, 0.5), loc="center left")
plt.grid()
ax = plt.gca()
```

```python
assert len(ax.lines) == 3
for l in ax.lines:
    ys = l.get_ydata()
    assert max(np.abs(ys)) < 1e-6
```

Comente.

<!-- grade_id: cell-3745cdedd0a58e23, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

## Questão 10: Interpolando em um intervalo maior

Agora, faça para a função `cauchy`, no intervalo $[-4,4]$.
Como a série de Taylor (em zero) é simplesmente a "fórmula da PG":

$$ \text{cauchy}(x) = 1 - x^2 + x^4 - x^6 \ldots, $$

que não converge fora de $(-1,1)$,
faça apenas dos polinômios de grau 5 (e depois 10) por Lagrange uniforme e por regressão.

<!-- grade_id: cell-33cb1e72bbe65fcd, points: 0, grade: false, solution: true -->
```python
# Ajuste em grau 5
# YOUR CODE HERE
raise NotImplementedError()

plt.legend(title="Método", bbox_to_anchor=(1, 0.5), loc="center left")
plt.grid()
ax = plt.gca()
```

<!-- grade_id: cell-a08adf82967bcfcb, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
for l in ax.lines:
    ys = l.get_ydata()
    assert max(np.abs(ys)) > 1e-1
```

<!-- grade_id: cell-722ec09e8834ab2e, points: 0, grade: false, solution: true -->
```python
# Ajuste em grau 10
# YOUR CODE HERE
raise NotImplementedError()

plt.legend(title="Método", bbox_to_anchor=(1, 0.5), loc="center left")
plt.grid()
ax = plt.gca()
```

<!-- grade_id: cell-b57e46fcd05b1481, points: 1, grade: true, solution: false -->
```python
assert len(ax.lines) == 2
for l in ax.lines:
    ys = l.get_ydata()
    assert max(np.abs(ys)) > 1e-2
```

O que você vê de diferente?

<!-- grade_id: cell-54c73cc651faa445, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

Faça, agora, para polinômios de grau 20. 

<!-- grade_id: cell-7eb0028d834ff5ff, points: 0, grade: false, solution: true -->
```python
# Ajuste em grau 20
# YOUR CODE HERE
raise NotImplementedError()
```

Como o erro de ajuste evolui, ao aumentar $n$, para o método de Lagrange?
E para o método de regressão?

<!-- grade_id: cell-ab3f6af99d5a9283, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE