![CC-BY-SA](https://mirrors.creativecommons.org/presskit/buttons/88x31/svg/by-sa.svg)


This notebook was created by [Bernardo Freitas Paulo da Costa](http://www.im.ufrj.br/bernardofpc),
and is licensed under Creative Commons BY-SA.

Antes de enviar este Teste, verifique que tudo está funcionando como esperado.
Por exemplo, **rode o código inteiro, do zero**.
Para isso, vá no menu, escolha _Kernel_, depois _Restart & Run All_.

Verifique, também, que você respondeu todas as questões:
* as questões de código têm `YOUR CODE HERE` (e você pode apagar o `raise NotImplemented` ao incluir sua resposta)
* as questões discursivas têm "YOUR ANSWER HERE".

---

```python
import numpy as np
```

# Teste 1: Bisseção

## Questão 1: Número de iterações da bisseção, e de chamadas à função `f`

Generalize o algoritmo da bisseção _iterativa_ para retornar,
- a (aproximação da) raiz,
- o número de bisseções (divisões por 2 do intervalo) feitas,
- o número de vezes que você chamou a função `f`

Use como critério de parada
- `xtol`, e
- `maxiter`,

ou seja, o algoritmo termina quando:
- seja possível garantir que o erro (absoluto) da resposta ("em $x$") seja menor do que `xtol`; OU
- quando o algoritmo já tiver feito `maxiter` itererações.

<!-- grade_id: cell-9c2f0bf1fe48da91, points: 0, grade: false, solution: true -->
```python
def bissecao(f, a, b, xtol=1e-8, maxiter=10):
    iters = 0
    calls = 0
    while abs(b-a) >= xtol and iters < maxiter:
        m = (a+b)/2
        if f(m)*f(a) < 0:
            b = m
        else:
            a = m
        iters += 1
        calls += 2
    return (a+b)/2, iters, calls
        
```

### 1.1: Testando com $\sqrt2$:

```python
def p2(x): return x**2 - 2
```

<!-- grade_id: cell-6216e0119f2a0e78, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, maxiter=100)
assert abs(raiz - np.sqrt(2)) < 1e-8
assert 25 <= niters <= 30
assert 25 <= ncalls
```

<!-- grade_id: cell-caee585cfc6b637f, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2, xtol=1e-3)
assert abs(raiz - np.sqrt(2)) < 1e-3
assert 9 <= niters <= 11
```

<!-- grade_id: cell-ed17d59cca42278d, points: 1, grade: true, solution: false -->
```python
raiz, niters, ncalls = bissecao(p2, 0, 2)
assert abs(raiz**2 - 2) > 1e-3
assert niters == 10
```

<!-- grade_id: cell-19379484bc5a3adf, points: 1, grade: true, solution: false -->
```python

```

### 1.2: Tolerância, número de iterações, respostas, ...

Observe o seguinte código:

```python
for err in np.logspace(-8, 0, num=9):
    r, ni, nc = bissecao(p2, 0, 2, xtol=err)
    print(f"xtol = {err:.1e} --> raiz = {r:.10f}, #iter = {ni:3d}")
```

Output:
```
xtol = 1.0e-08 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-07 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-06 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-05 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-04 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-03 --> raiz = 1.4150390625, #iter =  10
xtol = 1.0e-02 --> raiz = 1.4179687500, #iter =   8
xtol = 1.0e-01 --> raiz = 1.4062500000, #iter =   5
xtol = 1.0e+00 --> raiz = 1.2500000000, #iter =   2
```

Comente

<!-- grade_id: cell-856d87a71e8f8ec5, points: 2, grade: true, solution: true -->
A função *np.logspace()* retorna uma lista com 9 potências de 10 (padrão da função, quando não especificado) uniformemente espaçadas, em uma escala logarítmica, entre -8 (primeiro parametro) e 0 (segundo parâmetro). Dessa forma, o loop itera sobre essas potências e as usa para chamar, em cada uma das iterações, a função "bissecao", na qual a potência de uma dada iteração é passada como o parâmetro xtol. Isto é, a cada iteração, o erro mínimo que a função deve considerar corresponde à potência. No entanto, os resultados mostram que já na potência $10^{-3}$ a função atinge o número máximo padrão de iterações, não alterado neste algoritmo. Portanto, o menor erro possível alcançado é da ordem de $10^{-04}$, considerando a diferença entre o valor retornado para a raiz e o valor conhecido. Ou seja, independe do fato de o menor valor que este algoritmo passa, para o parâmetro de erro mínimo, xtol, seja na ordem de $10^{-08}$, já que o número de iterações máximo é atingido antes. Além disso, já que o menor erro passado corresponde ao erro padrão da função quando não é especificado um, se a ela for chamada sem ambos estes parâmetros, ela sempre irá terminar a bisseção, para a função p2 em questão, por ter primeiro atingido o máximo de iterações e não o erro mínimo.

Como seria o equivalente se o critério fosse o número de iterações?

<!-- grade_id: cell-c6d45f6ca0f539e4, points: 2, grade: true, solution: true -->
```python
for niters in np.linspace(10, 30, num=21):
    r, ni, nc = bissecao(p2, 0, 2, maxiter=niters)
    print(f"maxiter = {niters:.2f} --> raiz = {r:.10f}, #xtol = {r - np.sqrt(2):.2e}")
```

Output:
```
maxiter = 10.00 --> raiz = 1.4150390625, #xtol = 8.26e-04
maxiter = 11.00 --> raiz = 1.4145507812, #xtol = 3.37e-04
maxiter = 12.00 --> raiz = 1.4143066406, #xtol = 9.31e-05
maxiter = 13.00 --> raiz = 1.4141845703, #xtol = -2.90e-05
maxiter = 14.00 --> raiz = 1.4142456055, #xtol = 3.20e-05
maxiter = 15.00 --> raiz = 1.4142150879, #xtol = 1.53e-06
maxiter = 16.00 --> raiz = 1.4141998291, #xtol = -1.37e-05
maxiter = 17.00 --> raiz = 1.4142074585, #xtol = -6.10e-06
maxiter = 18.00 --> raiz = 1.4142112732, #xtol = -2.29e-06
maxiter = 19.00 --> raiz = 1.4142131805, #xtol = -3.82e-07
maxiter = 20.00 --> raiz = 1.4142141342, #xtol = 5.72e-07
maxiter = 21.00 --> raiz = 1.4142136574, #xtol = 9.50e-08
maxiter = 22.00 --> raiz = 1.4142134190, #xtol = -1.43e-07
maxiter = 23.00 --> raiz = 1.4142135382, #xtol = -2.42e-08
maxiter = 24.00 --> raiz = 1.4142135978, #xtol = 3.54e-08
maxiter = 25.00 --> raiz = 1.4142135680, #xtol = 5.60e-09
maxiter = 26.00 --> raiz = 1.4142135531, #xtol = -9.30e-09
maxite
```

Comente, também.

<!-- grade_id: cell-b7931432898e9741, points: 2, grade: true, solution: true -->
Desta vez, o algoritmo gera uma lista com 20 números entre 11 e 30, uniformemente espaçados, numa escala linear, dessa forma podemos evidenciar o máximo de iterações até que o critério de parada, xtol, atinja o seu valor padrão e faça a função parar.

### 1.3: Agora, um polinômio um pouco mais complicado:

```python
def p4(x): return x**4 - 3*x + 1
```

<!-- grade_id: cell-b8f933e77c525f69, points: 1, grade: true, solution: false -->
```python
r1, n1, c1 = bissecao(p4, 0, 1, maxiter=100)
assert p4(r1) < 1e-6
```

<!-- grade_id: cell-fab8ee9c8e55f544, points: 1, grade: true, solution: false -->
```python
r2, n2, c2 = bissecao(p4, 1, 2, maxiter=100)
assert n1 == n2
assert p4(r2) < 1e-6
```

<!-- grade_id: cell-faf0c1fe3a384747, points: 1, grade: true, solution: false -->
```python
r3, n3, c2 = bissecao(p4, 1, 3, maxiter=100)
assert n3 == n2+1
assert r3 == r2
```

<!-- grade_id: cell-1ec4a98add6ac67a, points: 1, grade: true, solution: false -->
```python

```

### 1.4: Relação entre `ncalls` e `niters`

Muito provavelmente, é possível calcular `ncalls` a partir de `niters`.
Qual a relação, para o seu programa, entre estes valores?
Porquê?

<!-- grade_id: cell-4a9c83444fc8ee6a, points: 1, grade: true, solution: true -->
```python
for iters in np.linspace(1, 100, num=10):
    r, ni, nc = bissecao(p2, 0, 2, xtol=1/(10**iters), maxiter=iters)
    print(f"ncalls = {nc}, niters = {ni}")
```

Output:
```
ncalls = 2, niters = 1
ncalls = 24, niters = 12
ncalls = 46, niters = 23
ncalls = 68, niters = 34
ncalls = 90, niters = 45
ncalls = 112, niters = 56
ncalls = 134, niters = 67
ncalls = 156, niters = 78
ncalls = 178, niters = 89
ncalls = 200, niters = 100
```

<!-- grade_id: cell-7d3444498035a363, points: 3, grade: true, solution: true -->
O algoritmo acima evidencia um padrão na função *bissecao()*: o número *ncalls* é o dobro do número *niters*. Isto porque, a cada iteração, a função bissecao usa a outra função argumento duas vezes, especificamente na linha escrito "if f(m)\*f(a)<0:", para verificar qual dos dois intervalos, [a,m] ou [m,b], contém uma raíz.  

## Questão 2: Bisseção em listas ordenadas

Se desejamos encontrar um número em uma lista, podemos percorrer cada um de seus índices
até achar o número, e retornamos o índice correspondente.

Mas, se a lista `l` possui números em ordem,
é possível usar uma variante da bisseção para encontrar este índice.
Observe que agora desejamos retornar um número **inteiro** (o índice!),
e que índices para listas sempre devem ser inteiros, portanto:
- cuidado ao escolher o "ponto médio"
- mas o critério de parada é "quando achar o índice".

### 2.1: Encontrando um número

Implemente a função `bissect_list(l, v)`, que retorna
- o índice `k` do elemento na lista `l`, que vale `v`; e
- o número de acessos à lista `l`.

Suponha que a lista `l` é não-decrescente (ou seja, `l[i] <= l[i+1]`).

<!-- grade_id: cell-8c9f49b0ba8b5d04, points: 0, grade: false, solution: true -->
```python
def bissect_list(l, v):
    na = 0 
    n = len(l)
    b = n-1
    a= 0
    while True:
        if n == 0 :
            return None, 0
        else:
            if n == 1:
                na+=1
                if l[a] == v:
                    return a, na
                else:
                    return None, na
            elif n == 2:
                na+=2
                if l[a] == v:
                    na-=1
                    return a, na
                elif l[b] == v:
                    return b, na
                else:
                    return None, na
            else:
                m = a + ((b-a)/2).__round__()
                lm = l[m]
                na += 1
                if lm>v:
                    b = m - 1
                    n = b - a + 1
                else:
                    if lm == v:
                        return m, na
                    else: 
                        a = m +1 
                        n = b - a + 1
```

<!-- grade_id: cell-d0e135df902b3955, points: 1, grade: true, solution: false -->
```python
assert bissect_list([], 42) == (None,0)
```

<!-- grade_id: cell-aa0a59b135a65159, points: 1, grade: true, solution: false -->
```python
l = [4]
assert bissect_list(l, 4) == (0,1)
```

<!-- grade_id: cell-ba03109833002639, points: 1, grade: true, solution: false -->
```python
assert bissect_list(l, 5) == (None,1)
```

<!-- grade_id: cell-84b5ba3a3384d223, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28]
for i,v in enumerate(l):
    idx, n = bissect_list(l, v)
    assert idx == i, f"v={v}, i={i}, idx={idx}"
    print(n)
```

Output:
```
3
2
3
1
3
2
3
```

Qual o maior número de acessos à lista que você precisou para encontrar um valor nela?
Este número depende da posição na lista?  Explique!

<!-- grade_id: cell-9059e7d02c8a0881, points: 2, grade: true, solution: true -->
O maior número de acessos foi 3. A posição influencia neste número, pois, para listas com tamanho maior do que 2, o primeiro valor verificado é o do meio, de indíce $m$, de determinado intervalo $[a,b]$  em uma iteração. Se um elemento não estiver nessa posição em nenhum dos intervalos das iterações, só será encontrado quando o intervalo for de dois elementos, quando a função verifica o valor nas extremidades, os dois únicos elementos. Se houver um número par de elementos, o índice $m$ se refere ao primeiro elemento da segunda metade do intervalo em consideração, isto porque o método especial __round__() é utilizado na divisão que determina $m$, que neste caso sempre resulta com 0,5 na casa decimal. 

### 2.2: Uma lista muito maior

```python
l = np.arange(-4000, 6000)
np.random.seed(1)
valores_testes = np.random.randint(low=-4000, high=6000, size=10)
for v in valores_testes:
    idx, n = bissect_list(l, v)
    assert l[idx] == v
    assert n <= 20
    print(f"v={v:5d}, idx={idx:5d}, n={n:3d}")
```

Output:
```
v=-3765, idx=  235, n= 13
v= 1192, idx= 5192, n= 13
v=-3095, idx=  905, n= 12
v= 3813, idx= 7813, n=  5
v=-1105, idx= 2895, n= 13
v= 1056, idx= 5056, n= 14
v=-3856, idx=  144, n= 12
v=  225, idx= 4225, n= 13
v= 3751, idx= 7751, n= 14
v= -538, idx= 3462, n= 11
```

Como o número de acessos varia, nesta outra lista?
E como ele se compara com o tamanho da lista?

<!-- grade_id: cell-1853e505f297b9b5, points: 2, grade: true, solution: true -->
Varia muito pouco, mas algumas vezes o número de acesso de um valor procurado fica bem menor que os dos outros, além de todos serem muito menor que o tamanho da lista, não passam de uma ordem de grandeza de $10^{1}$

### 3.3: Listas crescentes e decrescentes.

Altere o código anterior, para receber um argumento a mais, `decr`, que indica que a lista é **decrescente**.
Não reordene a lista!

<!-- grade_id: cell-e442fe23480d4b8d, points: 0, grade: false, solution: true -->
```python
def bissect(l, v, decr=False):
    na = 0 
    n = len(l)
    b = n-1
    a= 0
    while True:
        if n == 0 :
            return None, 0
        else:
            if n == 1:
                na+=1
                if l[a] == v:
                    return a, na
                else:
                    return None, na
            elif n == 2:
                na+=2
                if l[a] == v:
                    na-=1
                    return a, na
                elif l[b] == v:
                    return b, na
                else:
                    return None, na
            else:
                m = a + ((b-a)/2).__round__()
                lm = l[m]
                na += 1
                if lm>v:
                    if not decr:
                        b = m - 1
                    else:
                        a = m + 1
                    n = b - a + 1
                else:
                    if lm == v:
                        return m, na
                    else: 
                        if not decr:
                            a = m + 1
                        else :
                            b = m - 1
                        n = b - a + 1
```

<!-- grade_id: cell-8d0e1c185171dbce, points: 1, grade: true, solution: false -->
```python
l = [1,3,6,10,15,21,28,35,36]
assert bissect(l, 15)[0] == 4
```

<!-- grade_id: cell-75e5e557da5782e9, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 15, decr=True)[0] == 4
```

<!-- grade_id: cell-bd452b4939210a9a, points: 1, grade: true, solution: false -->
```python
l_rev = [36, 35, 28, 21, 15, 10, 6, 3, 1]
assert bissect(l_rev, 16, decr=True)[0] == None
```

<!-- grade_id: cell-4f806045de0b056c, points: 1, grade: true, solution: false -->
```python
l3 = [36, 35, 28, 21, 15, 10, 9, 6, 3, 2, 1]
assert bissect(l3, 15, decr=True)[0] == 4
```

### 2.4: Bônus

O número de acessos à lista para encontrar um número muda se esta está invertida?
Explique!

<!-- grade_id: cell-def3480b5f0dcea3, points: 2, grade: true, solution: true -->
Considerando os mesmos elementos ordenados diferentemente em duas listas, uma crescente e outra decrescente, o número de acesso geralmente será diferente - com exceção de alguns casos. Para todo intervalo de índice $[a,b]$ de tamanho maior que 2, o primeiro elemento verificado é o do índice $m$. Este índice ou é o elemento do meio, para um intervalo te tamanho impar, ou o primeiro da segunda metade, para um intervalo de tamanho par. Depois, quando o intervalo se reduz a dois elemento, é verificado o elemento num dado índice $a$, depois no $b$, primeiro e último respectivamente. Pelo fato do índice m de um intervalo par corresponder a esse elemento e pelo fato do elemento no índice $a$ ser verificado primeiro que o do $b$, alterando a ordem de crescente para decrescente, um valor pode mudar para um índice diferente. O exemplo mais simples para ilustrar isso é com uma lista de 2 elementos: [6,8]. Se quisermos encontrar 6, o número de acessos é 1, pois a lista tem tamanho 2 e o índice $a$ é o primeiro a ser verificado. Se mudarmos para decrescente, [8,6], o número de acessos é 2, pois primeiro verifica o do índice $a$, depois o do $b$. Numa lista com 4 elementos: [2,5,7,9], o elemento de índice $m$ é o 7, se estivermos procurando ele, usaríamos apenas 1 acesso. Alterando para ordem decrescente, [9,7,5,2], o 5 que está no índice $m$ agora, então 1 acesso seria feito e não teríamos encontrado o número 7. Será preciso considerar um novo intervalo, $[0,1]$, no qual o primeiro elemento será verificado e depois o segundo, onde está o valor procurado. Assim, o número total de acessos será 3. 